# PH 306 – Computational Physics
## Assignment Template

**Name:**

**Date:**

---

### Instructions
- Complete each problem in the cells marked **`# YOUR CODE HERE`**.
- Remove the `raise NotImplementedError()` line before writing your solution.
- Do **not** modify any read-only cells.
- Run all cells before submitting.

```python
import numpy as np
from scipy import integrate
```

---
## Problem 1 – Array Operations (2 points)

Create a NumPy array `x` containing 100 evenly spaced values from $0$ to $2\pi$ (inclusive),
then compute `y` as the element-wise sine of `x`.

<!-- grade_id: problem1-solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem1-tests, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(x, np.ndarray), "x must be a NumPy array"
assert len(x) == 100, "x must contain 100 elements"
assert np.isclose(x[0], 0.0), "x must start at 0"
assert np.isclose(x[-1], 2 * np.pi), "x must end at 2π"
assert isinstance(y, np.ndarray), "y must be a NumPy array"
assert np.allclose(y, np.sin(x)), "y must equal sin(x)"
### END HIDDEN TESTS
```

---
## Problem 2 – Numerical Integration (3 points)

Use `scipy.integrate.quad` to numerically integrate $f(x) = e^{-x^2}$ from $-\infty$ to $+\infty$.
Store the result (the integral value) in a variable named `result` and the estimated absolute
error in a variable named `error`.

*Hint:* The exact answer is $\sqrt{\pi}$.

<!-- grade_id: problem2-solution, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem2-tests, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert isinstance(result, float), "result must be a float"
assert isinstance(error, float), "error must be a float"
assert np.isclose(result, np.sqrt(np.pi), rtol=1e-5), (
    f"result={result} is not close to sqrt(π)={np.sqrt(np.pi)}"
)
### END HIDDEN TESTS
```