# PH 306 – Computational Physics
## Assignment Driver Notebook

**Name:**

**Date:**

---

### Instructions
- Complete all required functions in **`assignment.py`**.
- Use this notebook to run checks and inspect intermediate values.
- GitHub Classroom grading runs against **`assignment.py`**, not notebook cell output.
- Run all cells before submitting to verify your implementation.

```python
import numpy as np
import assignment
```

---
## Problem 1 – Array Operations (2 points)

In **`assignment.py`**, implement:

- `problem1_array_operations()`

It should return `(x, y)` where:

- `x` is 100 evenly spaced values from $0$ to $2\pi$ (inclusive)
- `y = sin(x)`

<!-- grade_id: problem1-solution, points: 0, grade: false, solution: true -->
```python
x, y = assignment.problem1_array_operations()
x[:5], y[:5]
```

<!-- grade_id: problem1-tests, points: 2, grade: true, solution: false -->
```python
### PUBLIC CHECKS (similar to autograder)
assert isinstance(x, np.ndarray), "x must be a NumPy array"
assert len(x) == 100, "x must contain 100 elements"
assert np.isclose(x[0], 0.0), "x must start at 0"
assert np.isclose(x[-1], 2 * np.pi), "x must end at 2π"
assert isinstance(y, np.ndarray), "y must be a NumPy array"
assert np.allclose(y, np.sin(x)), "y must equal sin(x)"
print("Problem 1 checks passed")
```

---
## Problem 2 – Numerical Integration (3 points)

In **`assignment.py`**, implement:

- `problem2_numerical_integration()`

It should return `(result, error)` where `result` is the numerical value of
$\int_{-\infty}^{+\infty} e^{-x^2} dx$ and `error` is SciPy's absolute error estimate.

*Hint:* The exact answer is $\sqrt{\pi}$.

<!-- grade_id: problem2-solution, points: 0, grade: false, solution: true -->
```python
result, error = assignment.problem2_numerical_integration()
result, error
```

<!-- grade_id: problem2-tests, points: 3, grade: true, solution: false -->
```python
### PUBLIC CHECKS (similar to autograder)
assert isinstance(result, float), "result must be a float"
assert isinstance(error, float), "error must be a float"
assert np.isclose(result, np.sqrt(np.pi), rtol=1e-5), (
    f"result={result} is not close to sqrt(π)={np.sqrt(np.pi)}"
)
print("Problem 2 checks passed")
```